# test 1
create a variable x and assign it the value 10

<!-- grade_id: cell-50c5e9654633, points: 0, grade: false, solution: true -->
```python

### BEGIN SOLUTION
x =10
### END SOLUTION
```

<!-- grade_id: cell-7f0029f4974f, points: 10, grade: true, solution: false -->
```python
#assertions
assert x == 10, "The value of x should be 10"
```

<!-- grade_id: cell-038ac8041934, points: 5, grade: true, solution: false -->
```python
assert x == 10, "The value of x should be 10"
```